---
title: 'Introduction'
description: 'We introduce how experiments help us find causal connections, and take a quick look at controlled experiments.'
---


## Course Trailer for Causal Inference with R - Experiments

```yaml
type: VideoExercise 
lang: r
xp: 50 
skills: 1
key: 4070242f87 
video_link: //player.vimeo.com/video/280255549  
```

---

## Introduction to Experiments

```yaml
type: VideoExercise 
lang: r
xp: 50 
skills: 1
key: 4070242f87 
video_link: //player.vimeo.com/video/205124502  
```

---

## Rainfall in Sonoma County and eBay sales

```yaml
type: MultipleChoiceExercise 
lang: r
xp: 50 
skills: 1
key: f571c52e35   
```


As mentioned in the video, rainfall in Sonoma County, California is highly correlated with eBay's total gross merchandise volume: 
![](https://assets.datacamp.com/production/repositories/1145/datasets/b97fef1a9b4ea4e8357ac2b39d04bf4c4e77cb34/ebayandsonomarainfall.png)

Which of the following answers could support a **causal explanation** for this relationship?


`@instructions`
- Across the world, people prefer to buy more eBay products on rainy days.
- Climate change has caused Sonoma County to have increased rainfall and has encouraged the global population to do more shopping from home.
- eBay sales and rainfall in Sonoma County have both been increasing at similar rates in recent years
- Rainfall motivates residents of Sonoma County to buy and sell goods on eBay at extraordinarily high rates

`@hint`


`@sct`

```{r}
msg1 = "This answer is close to correct, but not quite. Remember, eBay sells its goods to people across the world---why would the correlation be so strong with rainfall just in Sonoma County? Try again."
msg2 = "This is a possible explanation for the correlation between eBay sales and rainfall in Sonoma County, but not one that involves the rainfall causing the sales. Instead it would suggest that this correlation is spurious. Try again."
msg3 = "This is what it means for there to be a correlation between rainfall in Sonoma County and eBay sales, but it is not a causal explanation for this correlation. Try again."
msg4 = "Good Job. As ridiculous as this sounds, it is what a causal relationship between rainfall in Sonoma County and eBay sales might suggest. Of course, residents of Sonoma County would have to be providing a huge proportion of all eBay transactions for this to be correct."
test_mc(correct = 4, feedback_msgs = c(msg1,msg2,msg3,msg4))
```


---

## iPhone sales and eBay sales

```yaml
type: MultipleChoiceExercise 
lang: r
xp: 50 
skills: 1
key: caffaa946e   
```


Annual Apple iPhone sales are highly correlated with annual eBay sales. If this relationship is **spurious** (misleading and false), what might be causing this high correlation?


`@instructions`
- Most transactions on eBay are for Apple iPhones.
- Time causes technology sales and technology usage to increase
- Most people only shop on eBay using the iPhone's eBay app.
- Buying an iPhone makes people more comfortable with using technology for market transactions.

`@sct`

```{r}
msg1 = "This would suggest a causal relationship rather than a spurious correlation. Try again"
msg2 = "Correct! This would cause both iPhone and eBay sales to increase, but has nothing to say about a link between the two."
msg3 = "This would suggest a causal relationship rather than a spurious correlation. Try again"
msg4 = "This would suggest a causal relationship rather than a spurious correlation. Try again"
test_mc(correct = 2, feedback_msgs = c(msg1,msg2,msg3,msg4))
```


---

## A Quick Response

```yaml
type: VideoExercise 
lang: r
xp: 50 
skills: 1
key: 764f98ae81 
video_link: //player.vimeo.com/video/205124567  
```

---

## Controlled Experiments

```yaml
type: VideoExercise 
lang: r
xp: 50 
skills: 1
key: 4745db858e 
video_link: //player.vimeo.com/video/198212077  
```

---

## Buying a New Phone: Controlled experiments

```yaml
type: MultipleChoiceExercise 
lang: r
xp: 50 
skills: 1
key: ebe90c9c16   
```


Laurel is in the market for a new phone after cracking the screen on her old one. Her next phone choice will be based on which is more durable: the WePhone 10S or the Universe S10. She looks at a YouTube channel video where the host buys one of each phone and tests how much force is needed to crush each phone with a hydraulic press. If every WePhone or Universe from the factory is truly identical and if the phones are tested under exactly the same condition, then is just one crushing test enough to determine durability?


`@instructions`
- Yes
- No

`@sct`

```{r}
msg1 = "Correct, only two observations are needed to determine whether there is a causal effect between the treatment (phone type) and outcome (durability) if all experimental conditions between the phones are identical (i.e. If all other factors that could influence the outcome are controlled for). That's assuming you think a crush test is a valid one for determining real-world durability, of course!"
msg2 = "You might think not, but consider this as the most theoretically ideal case, and try again."
test_mc(correct = 1, feedback_msgs = c(msg1,msg2))
```


---

## Let’s Code: Comparing Ratios

```yaml
type: VideoExercise 
lang: r
xp: 50 
skills: 1
key: 75688f61ad 
video_link: //player.vimeo.com/video/279737801  
```

---

## Buying a New Phone: Comparing Ratios

```yaml
type: NormalExercise 
lang: r
xp: 100 
skills: 1
key: e157142b3a   
```


Laurel's video showed that the WePhone 10S was stronger than the Universe S10. Laurel is now interested in what factors caused the WePhone 10S to be more durable than the Universe S10. Laurel examines the model specifications of each phone, and assumes that whichever specification is most different between the phones is the cause of their different durabilities.


`@instructions`
- 1) Examine the provided phone specifications in dataframe`PhoneSpecifications` and identify the specification that has the `largest ratio` between the WePhone 10S and the Universe S10
- 2) Write out the `name` of that specification
- 3) Write out the `size` of this ratio (expressed as a quotient)

`@hint`
- For Solution1, make sure to identify the name of the specification precisely. A trick to doing this without writing it out is to use the `names` function, for example `names(PhoneSpecifications)[X]` where X equals the column of the specification
- For Solution2. The ratio of a specification between the WePhone 10S and Universe S10 can be calculated by dividing the specification value of the WePhone 10S by the specification value of the Universe S10. Make sure that your answer is not rounded. Use R's syntax to compute (for example, enter `Specification$X[1]/Specification$X[2]` where X is the specification.

`@pre_exercise_code`

```{r}
library(dplyr)
library(data.table)

PhoneSpecifications<-data.frame(Phone=c("WePhone 10S","UniverseS10"),Weight=c(190,180),Battery_Capacity=c(4000,3500),Diagonal_Length=c(5.5,5.2),Storage=c(64,64),Luminance=c(800,550))
```


`@sample_code`

```{r}
# 1) To begin, let's print the contents of the dataframe to the console by entering the dataframe name, PhoneSpecifications, into the print() function.

    Solution1<- print()

# 2) Now write the column name of the specification with the largest ratio between the WePhone 10S and the Universe S10 (Must be a character).

    Solution2<- ""

# 3) Now return the ratio between the WePhone10S and UniverseS10 of the phone specification you selected for Question 2 (do not round!).  
    
    Solution3<-

```


`@solution`

```{r}
Solution1<-print(PhoneSpecifications)

#This is how we solve this answer when the data are randomized, but an easier way to solve this questions is to look through each variable manually
    Ratios<-vector(length=5)
    for(i in 2:length(PhoneSpecifications[1,])){
      Ratios[i-1]<-PhoneSpecifications[1,i]/PhoneSpecifications[2,i]
    }
    Solution2<-names(PhoneSpecifications)[match(max(Ratios),Ratios)+1]

Solution3<-max(Ratios)
```


`@sct`

```{r}
test_object("Solution1")
test_object("Solution2")
test_object("Solution3")
success_msg("Good work! Keep in mind, even though there is a correlation between the treatment (luminance) and outcome (durability), this correlation is spurious; in reality, a phone's luminance (brightness) does not have a causal effect on its durability.
")
```


---

## Is it a Controlled Experiment?

```yaml
type: MultipleChoiceExercise 
lang: r
xp: 50 
skills: 1
key: a880db44ea   
```


As a staunch fiscal conservative, Peter is looking for ways to lower his daily expenses. Peter believes he might be able to lower his costs on car gasoline if he stops using his air conditioner and just opens his window. As an experiment, Peter tracks his gas usage during his commute to work over two days. During the first day, Peter uses his air conditioner. The second day, Peter opens his window. Is this a controlled experiment?


`@instructions`
- Yes
- No

`@sct`

```{r}
msg1 = "Can we be sure that that independent variable is the only factor being adjusted in Peter's experiment? Try again."
msg2 = "Correct! An experiment is considered a 'controlled' experiment when the independent variable is the only factor that is adjusted between the treatment and control. However, Peter's gas usage could be influenced by many factors that are outside of his control, like the traffic or the weather."
test_mc(correct = 2, feedback_msgs = c(msg1,msg2))
```

