--- 
title       : "Experiments 2: Controlled Experiments"
description : "These videos and exercises introduce controlled experiments."

--- type:VideoExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:4745db858e
## Controlled Experiments
*** =video_link
//player.vimeo.com/video/198212077

--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:ebe90c9c16
## Multiple Choice: Controlled experiments
Laurel is in the market for a new phone after cracking the screen on her old one. Her next phone choice will be based on which is more durable: the WePhone 10S or the Universe S10. She looks at a YouTube channel video where the host buys one of each phone and tests how much force is needed to crush each phone with a hydraulic press. If every WePhone or Universe from the factory is truly identical and if the phones are tested under exactly the same condition, then is just one crushing test enough to determine durability?

*** =instructions
- Yes
- No
*** =sct
```{r}
msg1 = "Correct, only two observations are needed to determine whether there is a causal effect between the treatment (phone type) and outcome (durability) if all experimental conditions between the phones are identical (i.e. If all other factors that could influence the outcome are controlled for). That's assuming you think a crush test is a valid one for determining real-world durability, of course!"
msg2 = "You might think not, but consider this as the most theoretically ideal case, and try again."
test_mc(correct = 1, feedback_msgs = c(msg1,msg2))
```

---
## Let’s Code: Comparing Ratios

```yaml
type: VideoExercise
lang: r
xp: 50
skills: 1
key: 86b5eff969
video_link: "//player.vimeo.com/video/279737801"
```

--- type:NormalExercise lang:r aspect_ratio:62.5 xp:50 skills:1 key:e157142b3a
## Comparing Ratios
Laurel's video showed that the WePhone 10S was stronger than the Universe S10. Laurel is now interested in what factors caused the WePhone 10S to be more durable than the Universe S10. Laurel examines the model specifications of each phone, and assumes that whichever specification is most different between the phones is the cause of their different durabilities.

*** =instructions
- 1) Examine the provided phone specifications in dataframe`PhoneSpecifications` and identify the specification that has the `largest ratio` between the WePhone 10S and the Universe S10
- 2) Write out the `name` of that specification
- 3) Write out the `size` of this ratio (expressed as a quotient)

*** =hint
- For Solution1, make sure to identify the name of the specification precisely. A trick to doing this without writing it out is to use the `names` function, for example `names(PhoneSpecifications)[X]` where X equals the column of the specification
- For Solution2. The ratio of a specification between the WePhone 10S and Universe S10 can be calculated by dividing the specification value of the WePhone 10S by the specification value of the Universe S10. Make sure that your answer is not rounded. Use R's syntax to compute (for example, enter `Specification$X[1]/Specification$X[2]` where X is the specification.


*** =pre_exercise_code
```{r}

library (dplyr)
library (data.table)
library (purr)

PhoneSpecifications<-data.frame(Phone=c("WePhone 10S","UniverseS10"),Weight=c(190,180),Battery_Capacity=c(4000,3500),Diagonal_Length=c(5.5,5.2),Storage=c(64,64),Luminance=c(800,550))
```
*** =sample_code
```{r}
# 1) To begin, let's print the contents of the dataframe to the console by entering the dataframe name, PhoneSpecifications, into the print() function.
#---- Question 1-------------------------------------#
    Solution1<- 
#----------------------------------------------------#

# 2) Now write the column name of the specification with the largest ratio between the WePhone 10S and the Universe S10 (Must be a character).
#---- Question 2-------------------------------------#
    Solution2<-    
#----------------------------------------------------#

# 3) Now return the ratio between the WePhone10S and UniverseS10 of the phone specification you selected for Question 2 (do not round!).  
#---- Question 3-------------------------------------#
    Solution3<- 
#----------------------------------------------------#
```
*** =solution
```{r}
Solution1<-print(PhoneSpecifications)

#loop through specification ratios to find max ratio
Ratios<-vector(length=5)
for(i in 2:length(PhoneSpecifications[1,])){
  Ratios[i-1]<-PhoneSpecifications[1,i]/PhoneSpecifications[2,i]
}
Solution2<-names(PhoneSpecifications)[match(max(Ratios),Ratios)+1]

Solution3<-max(Ratios)

```
*** =sct
```{r}

test_object("Solution1")
test_object("Solution2")
test_object("Solution3")
test_function("print", incorrect_msg = "Did you use the `print` function?")
test_error()
success_msg("Good work! Keep in mind, even though there is a correlation between the treatment (luminance) and outcome (durability), this correlation is spurious; in reality, a phone's luminance (brightness) does not have a causal effect on its durability.
")
```


--- type:MultipleChoiceExercise lang:r xp:50 skills:1 key:a880db44ea
## Multiple Choice: Is it a Controlled Experiment?
As a staunch fiscal conservative, Peter is looking for ways to lower his daily expenses. Peter believes he might be able to lower his costs on car gasoline if he stops using his air conditioner and just opens his window. As an experiment, Peter tracks his gas usage during his commute to work over two days. During the first day, Peter uses his air conditioner. The second day, Peter opens his window. Is this a controlled experiment?

*** =instructions
- Yes
- No
*** =sct
```{r}
msg1 = "Can we be sure that that independent variable is the only factor being adjusted in Peter's experiement? Try again."
msg2 = "Correct! An experiment is considered a 'controlled' experiment when the independent variable is the only factor that is adjusted between the treatment and control. However, Peter's gas usage could be influenced by many factors that are outside of his control, like the traffic or the weather."
test_mc(correct = 2, feedback_msgs = c(msg1,msg2))
```
